<a href="https://colab.research.google.com/github/ellahad/Phase-2-Code-Challenge/blob/main/index.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Phase 2 Code Challenge

This code challenge is designed to test your understanding of the Phase 2 material. It covers:

- Normal Distribution
- Statistical Tests
- Bayesian Statistics
- Linear Regression

_Read the instructions carefully_. You will be asked both to write code and to answer short answer questions.

## Code Tests

We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. To determine what the issue is, read the comments in the code test cells, the error message you receive, and the item instructions.

## Short Answer Questions

For the short answer questions...

* _Use your own words_. It is OK to refer to outside resources when crafting your response, but _do not copy text from another source_.

* _Communicate clearly_. We are not grading your writing skills, but you can only receive full credit if your teacher is able to fully understand your response.

* _Be concise_. You should be able to answer most short answer questions in a sentence or two. Writing unnecessarily long answers increases the risk of you being unclear or saying something incorrect.

```python
# Run this cell without changes to import the necessary libraries

import itertools
import numpy as np
import pandas as pd
from numbers import Number
from scipy import stats
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

import pickle
```

---
## Part 1: Normal Distribution [Suggested time: 20 minutes]
---
In this part, you will analyze check totals at a TexMex restaurant. We know that the population distribution of check totals for the TexMex restaurant is normally distributed with a mean of \\$20 and a standard deviation of \\$3.

### 1.1) Create a numeric variable `z_score_26` containing the z-score for a \\$26 check.

**Starter Code**

    z_score_26 =

<!-- grade_id: cell-b947773fd0a27e5d, points: 0, grade: false, solution: true -->
```python
#standad deviation calculation
#assign values
x = 26
mean = 20
std_dev = 3
z_score_26 = (x - mean) / std_dev
print(z_score_26)
```

Output:
```
2.0
```

<!-- grade_id: cell-a72f04fffcbada78, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named z_score_26

assert isinstance(z_score_26, Number)

```

### 1.2) Create a numeric variable `p_under_26` containing the approximate proportion of all checks that are less than \\$26.

Hint: Use the answer from the previous question along with the empirical rule, a Python function, or this [z-table](https://www.math.arizona.edu/~rsims/ma464/standardnormaltable.pdf).

**Starter Code**

    p_under_26 =

<!-- grade_id: cell-a223448d6f75a998, points: 0, grade: false, solution: true -->
```python
#finding mean and std dev
mean_check = 50
std_dev_check = 15
check_value = 26
# Find the estimated percentage
p_under_26 = stats.norm.cdf(check_value, loc=mean_check, scale=std_dev_check)
f"The rough percentage of checks under $26: {round(p_under_26, 4)}"
```

Output:
```
'The rough percentage of checks under $26: 0.0548'
```

<!-- grade_id: cell-626d2e6dea5a86ac, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_under_26

assert isinstance(p_under_26, Number)

# These tests confirm that p_under_26 is a value between 0 and 1

assert p_under_26 >= 0
assert p_under_26 <= 1

```

### 1.3) Create numeric variables `conf_low` and `conf_high` containing the lower and upper bounds (respectively) of a 95% confidence interval for the mean of one waiter's check amounts using the information below.

One week, a waiter gets 100 checks with a mean of \\$19 and a standard deviation of \\$3.

**Starter Code**

    n = 100
    mean = 19
    std = 3
    
    conf_low =
    conf_high =

<!-- grade_id: cell-fb0a963f092cb64d, points: 0, grade: false, solution: true -->
```python
#define values

mean = 19
sample_size = 100
std_dev = 3
confidence_level = 0.95
# margin error
margin_error = 1.96 * (std_dev / (sample_size ** 0.5))
#Find the confidence interval.
conf_low = mean - margin_error
conf_high = mean + margin_error
f"95% Confidence Interval: {round(conf_low, 2)} - {round(conf_high, 2)}"
```

Output:
```
'95% Confidence Interval: 18.41 - 19.59'
```

<!-- grade_id: cell-d995ba92d5f4f1f3, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created numeric variables named conf_low and conf_high

assert isinstance(conf_low, Number)
assert isinstance(conf_high, Number)

# This test confirms that conf_low is below conf_high

assert conf_low < conf_high

# These statements print your answers for reference to help answer the next question

print('The lower bound of the 95% confidence interval is {}'.format(conf_low))
print('The upper bound of the 95% confidence interval is {}'.format(conf_high))

```

Output:
```
The lower bound of the 95% confidence interval is 18.412
The upper bound of the 95% confidence interval is 19.588
```

### 1.4) Short Answer: Interpret the 95% confidence interval you just calculated in Question 1.3.

<!-- grade_id: cell-8212712ac7485fa9, points: 1, grade: true, solution: true -->

> I am quite sure, with 95% confidence, that the actual average amount for one waiter's checks is somewhere between 18.44 and 19.56, according to the provided sample data.








---
## Part 2: Statistical Testing [Suggested time: 20 minutes]
---
The TexMex restaurant recently introduced queso to its menu.

We have a random sample containing 2000 check totals, all from different customers: 1000 check totals for orders without queso ("no queso") and 1000 check totals for orders with queso ("queso").

In the cell below, we load the sample data for you into the arrays `no_queso` and `queso` for the "no queso" and "queso" order check totals, respectively.

```python
# Run this cell without changes

# Load the sample data
no_queso = pickle.load(open('no_queso.pkl', 'rb'))
queso = pickle.load(open('queso.pkl', 'rb'))
```

Output:
```
FileNotFoundError                         Traceback (most recent call last)
<ipython-input-65-7e381d913ece> in <cell line: 4>()
      2 
      3 # Load the sample data
----> 4 no_queso = pickle.load(open('no_queso.pkl', 'rb'))
      5 queso = pickle.load(open('queso.pkl', 'rb'))

FileNotFoundError: [Errno 2] No such file or directory: 'no_queso.pkl'
```

### 2.1) Short Answer: State null and alternative hypotheses to use for testing whether customers who order queso spend different amounts of money from customers who do not order queso.

<!-- grade_id: cell-1992605f4c2b8751, points: 1, grade: true, solution: true -->

> For null hypothesis there is no spending difference between customers who order queso and those who do not while alternative hypothesis customers who order queso spend a different amount compared to those who do not.

### 2.2) Short Answer: What would it mean to make a Type I error for this specific hypothesis test?

Your answer should be _specific to this context,_  not a general statement of what Type I error is.

<!-- grade_id: cell-f5e449677777b2fb, points: 1, grade: true, solution: true -->

> Making a Type I error means mistakenly assuming a spending difference between
customers who order queso and those who don't, even if there isn't actually any difference.




### 2.3) Create a numeric variable `p_value` containing the p-value associated with a statistical test of your hypotheses.

You must identify and implement the correct statistical test for this scenario. You can assume the two samples have equal variances.

Hint: Use `scipy.stats` to calculate the answer - it has already been imported as `stats`. Relevant documentation can be found [here](https://docs.scipy.org/doc/scipy/reference/stats.html#statistical-tests).

**Starter Code**

    p_value =

<!-- grade_id: cell-c2e90acfaebc8303, points: 0, grade: false, solution: true -->
```python
p_value = ttest_ind(queso, no_queso, equal_var=True)[1]

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-66-7ee342158e1d> in <cell line: 1>()
----> 1 p_value = ttest_ind(queso, no_queso, equal_var=True)[1]

NameError: name 'queso' is not defined
```

<!-- grade_id: cell-2364f38312cbc15e, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created a numeric variable named p_value

assert isinstance(p_value, Number)

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-67-5865ab15e3b2> in <cell line: 3>()
      1 # These tests confirm that you have created a numeric variable named p_value
      2 
----> 3 assert isinstance(p_value, Number)

NameError: name 'p_value' is not defined
```

### 2.4) Short Answer: Can you reject the null hypothesis using a significance level of $\alpha$ = 0.05? Explain why or why not.

<!-- grade_id: cell-92ea8835a9a326eb, points: 1, grade: true, solution: true -->
> To decide if I can reject the null hypothesis at a significance level of 0.05, I would check if the p-value is smaller than 0.05.
> If the p-value is below 0.05, this means that i have evidence to reject the null hypothesis.
> If the p-value is above 0.05, ythis means that i lack evidence to reject the null hypothesis.



---
## Part 3: Bayesian Statistics [Suggested time: 15 minutes]
---
A medical test is designed to diagnose a certain disease. The test has a false positive rate of 10%, meaning that 10% of people without the disease will get a positive test result. The test has a false negative rate of 2%, meaning that 2% of people with the disease will get a negative result. Only 1% of the population has this disease.

### 3.1) Create a numeric variable `p_pos_test` containing the probability of a person receiving a positive test result.

Assume that the person being tested is randomly selected from the broader population.

**Starter Code**
    
    false_pos_rate = 0.1
    false_neg_rate = 0.02
    population_rate = 0.01
    
    p_pos_test =

<!-- grade_id: cell-30ff893c7c30f29d, points: 0, grade: false, solution: true -->
```python
# Starter Code
false_pos_rate = 0.1
false_neg_rate = 0.02
population_rate = 0.01

# calculating the probability of a person receiving a positive test result
p_pos_test = population_rate * (1 - false_neg_rate) + (1 - population_rate) * false_pos_rate
```

<!-- grade_id: cell-2d49a1068c9e46b2, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_pos_test

assert isinstance(p_pos_test, Number)

# These tests confirm that p_pos_test is a value between 0 and 1

assert p_pos_test >= 0
assert p_pos_test <= 1

```

### 3.2) Create a numeric variable `p_pos_pos` containing the probability of a person actually having the disease if they receive a positive test result.

Assume that the person being tested is randomly selected from the broader population.

Hint: Use your answer to the previous question to help answer this one.

**Starter Code**
    
    false_pos_rate = 0.1
    false_neg_rate = 0.02
    population_rate = 0.01
    
    p_pos_pos =

<!-- grade_id: cell-742273642f14d1ab, points: 0, grade: false, solution: true -->
```python
#Starter Code

false_pos_rate = 0.1
false_neg_rate = 0.02
population_rate = 0.01

p_pos_pos =(1 - false_neg_rate) * population_rate / p_pos_test
```

<!-- grade_id: cell-4825f55ab11b5fff, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_pos_pos

assert isinstance(p_pos_pos, Number)

# These tests confirm that p_pos_pos is a value between 0 and 1

assert p_pos_pos >= 0
assert p_pos_pos <= 1

```

---

## Part 4: Linear Regression [Suggested Time: 20 min]
---

In this section, you'll run regression models with advertising data. The dataset includes the advertising spending and sales for 200 products. These products are a random sample of 200 products from a larger population of products.

The dataset has three features - `TV`, `radio`, and `newspaper` - that describe how many thousands of advertising dollars were spent promoting the product. The target, `sales`, describes how many millions of dollars in sales the product had.

```python
# Run this cell without changes

import statsmodels
from statsmodels.formula.api import ols

data = pd.read_csv('./advertising.csv').drop('Unnamed: 0', axis=1)
X = data.drop('sales', axis=1)
y = data['sales']
```

Output:
```
/usr/local/lib/python3.10/dist-packages/pandas/io/parsers/readers.py in _make_engine(self, f, engine)
   1733                 if "b" not in mode:
   1734                     mode += "b"
-> 1735             self.handles = get_handle(
   1736                 f,
   1737                 mode,

/usr/local/lib/python3.10/dist-packages/pandas/io/common.py in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    854         if ioargs.encoding and "b" not in ioargs.mode:
    855             # Encoding
--> 856             handle = open(
    857                 handle,
    858                 ioargs.mode,

FileNotFoundError: [Errno 2] No such file or directory: './advertising.csv'
```

```python
from google.colab import drive
drive.mount('/content/drive')
```

### 4.1) Create a variable `tv_mod_summ` containing the summary of a fit StatsModels `ols` model for a linear regression using `TV` to predict `sales`.

**Starter Code**

    tv_mod = ols(
    tv_mod_summ = tv_mod.summary()

<!-- grade_id: cell-86323454c9b2aa75, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84e163360346b9f6, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a variable named tv_mod_summ containing a StatsModels Summary

assert type(tv_mod_summ) == statsmodels.iolib.summary.Summary

# This statement prints your answer for reference to help answer the next question

print(tv_mod_summ)

```

### 4.2) Short Answer: Is there a statistically significant relationship between TV advertising spend and sales in this model? How did you determine this from the model output?

This question is asking you to use your findings from the sample in your dataset to make an inference about the relationship between TV advertising spend and sales in the broader population.

<!-- grade_id: cell-ebd4cb97464e0212, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### 4.3) Short Answer: Run the cell below to produce a correlation matrix. Given the output, would you expect any collinearity issues if you included all of these features in one regression model?

```python
# Run this cell

X.corr()
```

<!-- grade_id: cell-c1e98a26b16748b1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### 4.4) Create a variable `all_mod_summ` containing the summary of a fit StatsModels `ols` model for a multiple regression using `TV`, `radio`, and `newspaper` to predict `sales`.

**Starter Code**

    all_mod = ols(
    all_mod_summ = all_mod.summary()

<!-- grade_id: cell-a9cc84d7d7666b0f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1f7f745c780b07b6, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a variable named all_mod_summ containing a StatsModels Summary

assert type(all_mod_summ) == statsmodels.iolib.summary.Summary

# This statement prints your answer for reference to help answer the next question

print(all_mod_summ)

```

### 4.5) Short Answer: Does this model do a better job of explaining sales than the previous model using only the `TV` feature? Explain how you determined this based on the model output.

<!-- grade_id: cell-7e27f10fcaaed6a9, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE